# Week 4 HW part 1:  loading NASA weblogs into RDDs

During lecture we hand-generated tiny amounts of data (e.g. `[5, 3, 65, 2]`) on the driver (this notebook) and then pushed that data out to the workers using `parallelize()`.  This created an RDD.

In this homework we will create an RDD in a different way:  the workers will read data by themselves from NASA weblog files that you already downloaded to the local filesystem.

```python
from pyspark import SparkContext
sc = SparkContext('local', 'weblog_analysis')
```

Using the `textFile()` function, we can tell spark to read files into an RDD (each line becomes a separate element in the RDD).

Remember that the NASA weblogs are compressed (gzipped).  Spark is smart enough to detect this and decompress automatically for us.

Also, using a wildcard pattern `*`, we can read in many files at once.

```python
logs_rdd = sc.textFile('NASA_access_log_*.gz')
```

Let's have a look at a few of the records

```python
logs_rdd.take(5)
```

```python
# We didn't specify the number of partitions, so let's see what Spark defaults to
logs_rdd.getNumPartitions()
```

Write code to determine how many log lines are in the files altogether.

<!-- grade_id: cell-33cb0bda4bf47600, points: 0, grade: false, solution: true -->
```python
# How many lines are in this file?  Do you think this is this an ACTION, or a TRANSFORMATION?
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5e055896f3afa0da, points: 10, grade: true, solution: false -->
```python
assert num_lines == 3461613
```

Write code that creates a new RDD containing the *lengths* of the lines.  For example, if `logs_rdd` looked like this:
```
['this',
 'is',
 'getting',
 'more',
 'challenging',
 ...]
```
then the new RDD `line_lengths_rdd` would look like this:
```
[4,
 2,
 7,
 4,
 11,
 ...]
```

<!-- grade_id: cell-209a46244a57b8c0, points: 0, grade: false, solution: true -->
```python
# NOTE: the assert that tests this cell is a few cells down
# YOUR CODE HERE
raise NotImplementedError()
```

We will reuse `line_lengths_rdd` again, so let's instruct Spark to cache (remember) it.  This way it doesn't need to be recomputed in the future.

If we don't do this then, after the *action* is taken, there is no guarantee that Spark will bother to remember the RDD that you computed (Spark starts cleaning up to free up some memory for future computations).

```python
line_lengths_rdd.persist()
```

<!-- grade_id: cell-a2d032710aac18d1, points: 10, grade: true, solution: false -->
```python
assert line_lengths_rdd.take(5) == [122, 71, 96, 98, 95]
```

Use `.reduce()` to compute the total number of characters in the dataset

<!-- grade_id: cell-0ee614360fdcf263, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2c3fd8525d877350, points: 10, grade: true, solution: false -->
```python
assert total_chars == 369594520
```

### Plotting line lengths

It is time to learn a little about graphing.  There are several popular graphing libraries for Python.  The oldest, and still most used, is `matplotlib`.  A tutorial can be found here:

https://pythonprogramming.net/bar-chart-histogram-matplotlib-tutorial/

Below we will plot a histogram of line lengths.  Admittedly this isn't very useful, but we're not ready to parse the logs to extract more interesting insights yet (that will be your Exam 1).

Make sure that your histogram looks reasonably professional - it needs to have labeled axes and a title.

But first, we need to subsample!

#### Subsampling

We saw above that there are over 3 million lines of logs.  For the purposes of plotting we should subsample
using the `.sample()` transformation.

This tutorial describes some of the most common transformations and actions, including `.sample()`:
    
https://www.analyticsvidhya.com/blog/2016/10/using-pyspark-to-perform-transformations-and-actions-on-rdd/

Collect a sample (**with replacement**) of about 5000 line lengths into a new RDD:

<!-- grade_id: cell-d51cf61030c53e85, points: 0, grade: false, solution: true -->
```python
sample_fraction = 5000/num_lines

seed = 2343243  # I just randomly typed that junk

# We will fixed the seed so that you get the same random sample that I get.

# EXPLANATION:  there is no such thing as a true random number generator.  
# Algorithms are deterministic.
# However, if you know enough about Number Theory then you can design an algorithm 
# that "looks" random.
# These types of algorithms are called "pseudorandom".

# All pseudorandom algorithms need to start somewhere (an initial number, or "seed").
# Given the same seed, the pseudorandom algorithm will always produce the same sequence
# of "random" numbers (no matter how many times it is run).

# In a real analysis the seed should be as random as possible, and is automatically 
# gathered in the beginning by the algorithm using random input from the
# outside world (e.g. last 10 keystrokes, current timestamp, etc).



# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d76d2352455697be, points: 10, grade: true, solution: false -->
```python
assert sample_line_lengths_rdd.take(5) == [86, 135, 104, 124, 103]
```

Now use a `.collect()` action to pull the 5000 samples back to the driver and plot them in a histogram:

```python
line_lengths = sample_line_lengths_rdd.collect()
```

<!-- grade_id: cell-b7596f1711a96468, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt
import seaborn  # seaborn just makes plots look a little nicer

# plot `line_lengths` in a histogram here
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```